# Analiza danych przestrzennych - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają na celu wyłapać podstawowe błędy. Przejście przez kod wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem). Rozmiar poziomy figur nie powinien przekraczać 20 cali.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu: IAD12
Członkowie: Izabela Karczewska 409957, Natalia Kołodziejczyk 409695, Natalia Kubańska 411933

---

## Zestaw zadań 2: Elementy statystyki opisowej 2 (Notatnik 1/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Przygotowanie danych [10 pkt]

Pliki raw_data_1.csv, raw_data_2.csv i raw_data_3.xlsx zawierają dane pobrane ze stron:</br>
https://databank.worldbank.org/,</br>
https://unstats.un.org/UNSDWebsite/,</br>
https://datahelpdesk.worldbank.org/knowledgebase/articles/906519-world-bank-country-and-lending-groups.

Zaimportuj zawarte w plikach dane do notatnika i utwórz na ich podstawie tablicę `recession_data_1` wykonaną według następujących wytycznych:

- Tablica powinna zawierać wyłącznie następujące elementy:</br>
    Indeks: Country Code</br>
        Wiersze powinny być indeksowane trzyliterowymi kodami krajów.</br>
    Kolumny: Country Name, Region Name, Income Group, 2006, 2007, 2008, 2009, 2010, 2011 (w tej kolejności, lata w nagłówkach jako zmienne typu str)</br>
        Kolumna Income Group zawiera informację o stopniu zamożności państwa wg danych z 2008 roku.
        Kolumny 2006-2011 zawierają wartości wskaźnika GDP growth (annual %) dla poszczególnych lat.
- Zmień występujące w kolumnie Income Category kody (H, UM, LM, L) na pełne wyrażenia (High, Upper middle, Lower middle, Low).
- Tablica nie powinna zawierać wierszy, w których znajdują się wartości NaN (w dowolnej kolumnie) oraz wiersza dotyczącego Libii.

Zapisz tablicę do pliku recession_data_1.pkl.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było po 5 pierwszych i ostatnich wierszy, a zawarte w tablicy liczby (za wyjątkiem lat w nagłówkach kolumn) wyświetlały się z dokładnością do 2 miejsca po przecinku.

<!-- grade_id: cell-aa85199b5cb633d2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df1 = pd.read_csv("raw_data_1.csv",index_col="Country Code",skipfooter=7,engine='python')
df1.replace("..",np.NaN,inplace=True)
df2 = pd.read_csv("raw_data_2.csv",delimiter=";",index_col="ISO-alpha3 Code")
df3 = pd.read_excel("raw_data_3.xlsx",sheet_name="Country Analytical History",skiprows = [*range(0,5),*range(6,11),*range(229,232)],skipfooter=2,index_col=0)
df3.replace("..",np.NaN,inplace=True)
display(df1)
```

Output:
```
Series Name        Series Code           Country Name  \
Country Code                                                                    
AFG           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG            Afghanistan   
ALB           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG                Albania   
DZA           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG                Algeria   
ASM           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG         American Samoa   
AND           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG                Andorra   
...                             ...                ...                    ...   
VIR           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG  Virgin Islands (U.S.)   
PSE           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG     West Bank and Gaza   
YEM           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG            Yemen, Rep.   
ZMB           GDP growth (annual %)  NY.GDP.MKTP.KD.ZG                 Zambia   
ZWE           GDP growth (annual %)  NY.GDP.MKTP.KD.
```

```python
df1.rename(columns = {col : col.split(" ")[0] for col in df1.columns[3:]},inplace = True)
df1 = df1[["Country Name", "2006", "2007","2008","2009","2010","2011"]]
df1 = df1[df1["Country Name"] != "Libya"]
df1.dropna(inplace=True)
print(len(df1))
```

Output:
```
201
```

```python
df2 = df2[["Region Name"]]
df2.dropna(inplace=True)
display(df2)
```

Output:
```
Region Name
ISO-alpha3 Code            
DZA                  Africa
EGY                  Africa
LBY                  Africa
MAR                  Africa
SDN                  Africa
...                     ...
WSM                 Oceania
TKL                 Oceania
TON                 Oceania
TUV                 Oceania
WLF                 Oceania

[248 rows x 1 columns]
```

```python
df3 = df3[[2008]]
df3.dropna(inplace=True)
display(df3)
```

Output:
```
2008
AFG    L
ALB   LM
DZA   UM
ASM   UM
AND    H
..   ...
YEM    L
ZMB    L
ZWE    L
MYT   UM
ANT    H

[211 rows x 1 columns]
```

```python
recession_data_1 = df1.join(df2, how = "inner").join(df3, how = "inner")
recession_data_1.replace({"L":"Low","LM":"Lower middle","H":"High","UM":"Upper middle"},inplace=True)
recession_data_1.rename(columns={2008:"Income Group"},inplace=True)
recession_data_1=recession_data_1[["Country Name", "Region Name", "Income Group", "2006", "2007", "2008", "2009", "2010", "2011"]]
recession_data_1.index.name="Country Code"
#display(recession_data_1)
recession_data_1[['2006', '2007', '2008', '2009', '2010', '2011']] = recession_data_1[['2006', '2007', '2008', '2009', '2010', '2011']].apply(pd.to_numeric)
```

```python
recession_data_1.iloc[[*range(5), *range(-5, 0)]].style.format({'2006':'{:.2f}', '2007':'{:.2f}', '2008':'{:.2f}', 
                                                                '2009':'{:.2f}', '2010':'{:.2f}', '2011':'{:.2f}'})
```

Output:
```
<pandas.io.formats.style.Styler at 0x20c7f943b50>
```

```python
recession_data_1.to_pickle('recession_data_1.pkl')
```

<!-- grade_id: cell-683b0a3c6c74d0c6, points: 10, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert recession_data_1.shape == (198, 9)
# Test 2 Kolumny
assert len(set(list(recession_data_1.columns)) & set(['Country Name', 'Region Name', 'Income Group', '2006', '2007', '2008', '2009', '2010', '2011'])) == 9
# Test 3 Indeks
assert (recession_data_1.index.name == 'Country Code') == True
assert recession_data_1.index.values.min() == 'ABW'
assert recession_data_1.index.values.max() == 'ZWE'
assert np.all(recession_data_1.index.values != 'LBY') == True
# Test 4 Zawartość kolumn tekstowych
assert np.all([set(recession_data_1["Region Name"].unique()) == set(['Asia', 'Europe', 'Africa', 'Oceania', 'Americas'])])
assert np.all([recession_data_1["Region Name"].value_counts().values == [50, 49, 43, 40, 16]])
assert np.all([set(recession_data_1["Income Group"].unique()) == set(['Low', 'Lower middle', 'Upper middle', 'High'])])
assert np.all([recession_data_1["Income Group"].value_counts().values == [60, 53, 44, 41]])
# Test 5 Zawartość kolumn liczbowych
assert np.all(np.isclose(recession_data_1[['2006', '2007', '2008', '2009', '2010', '2011']].min().values, np.array([ -6.87146322, -6.68402778, -17.66894633,
                                                                                                                    -17.57322176, -8.92417589, -12.7148969 ]))) == True
assert np.all(np.isclose(recession_data_1[['2006', '2007', '2008', '2009', '2010', '2011']].max().values, np.array([34.5 , 24.99999985, 17.79910913, 21.39052839,
                                                                                                                    25.12275964, 21.61652888]))) == True
```